<img src="https://assets.ensam.eu/logo/fr/logo-trans-322x84.png" width="256px" align="left"/>

Avant de me rendre votre notebook, assurez-vous que tout tourne comme attendu.
Pour cela, **redémarrer le kernel** (dans la barre de menu, selectionnez *Kernel $\rightarrow$ Restart*) et ensuite **exécutez toutes les cellules** (dans la barre de menu, selectionnez *Cellules $\rightarrow$ Exécuter toutes les cellules*).

Assurez-vous par ailleurs d'avoir écrit votre code partout où il est dit `YOUR CODE HERE` ainsi que les réponses libres où il est écrit "YOUR ANSWER HERE".
Enfin, rentrez votre nom et celui des personnes avec qui vous avez travaillés ci-dessous.

```python
NAME = ""
COLLABORATORS = ""
```

---

# TP n°1 : Introduction à Python

---

<img src="https://upload.wikimedia.org/wikipedia/commons/thumb/f/f8/Python_logo_and_wordmark.svg/1200px-Python_logo_and_wordmark.svg.png" width="256px" align="left"/>

L'objectif de ce tout premier TP est de vous familiariser avec la syntaxe de base en `Python`.
`Python` est un langage de programmation crée en 1991 par le Néerlandais [Guido van Rossum](https://fr.wikipedia.org/wiki/Guido_van_Rossum).
Depuis sa création, le langage est régulièrement amélioré, notamment grâce au travail de la [Python Software Foundation](https://fr.wikipedia.org/wiki/Python_Software_Foundation), et a été placé sous une [licence libre](https://fr.wikipedia.org/wiki/Licence_libre) proche de la [licence BSD](https://fr.wikipedia.org/wiki/Licence_BSD).

`Python` fonctionne sur la plupart des plates-formes informatiques, des smartphones jusqu'aux supers calculateurs.
Il est par ailleurs disponible pour la plupart des systèmes d'exploitation, de Windows à Linux en passant par macOS, Android et iOS.
Cette grande disponibilité et son caractère libre font qu'aujourd'hui `Python` fait parti des langages de programmation les plus utilisés dans le monde.
Par ailleurs, sa syntaxe relativement simple à prendre en main est également l'une des raisons pour lesquelles `Python` est souvent le premier langage de programmation enseigné aux étudiants.

---

**Installation** &emsp; Dans l'éventualité où vous n'auriez pas accès aux ordinateurs de l'école, il existe de nombreuses façons d'installer `Python` sur votre ordinateur personnel.

<img src="https://firefortysix.files.wordpress.com/2020/08/anaconda-logo-1.jpg?w=1120" width="256px" align="left"/>

La méthode la plus simple et la plus complète, peu importe votre OS, reste néanmoins d'installer la distribution `Python` proposée par [Anaconda](https://www.anaconda.com/products/individual).
Il s'agit d'une distribution `Python` extrêmement complète utilisée par plus de 25 millions d'utilisateurs à travers le monde, notamment par l'école.
L'installation est très simple et vous garantira d'avoir absolument tous les outils `Python` dont vous pourriez avoir besoin, aussi bien dans le cadre de ce cours que plus généralement.
En particulier, [Anaconda](https://www.anaconda.com/products/individual) installera automatiquement [Jupyter](https://jupyter.org/) vous permettant ainsi d'utiliser les **notebooks** sur lesquels ce cours est basé.

<img src="https://colab.research.google.com/img/colab_favicon_256px.png" width="128px" align="left"/>

Par ailleurs, si il est vous est impossible d'installer directement Anaconda car l'ordinateur portable que vous utilisez appartient à votre entreprise et que l'accord du service informatique est nécessaire, notez que Google met gratuitement à disposition ses capacités de calcul pour utiliser en ligne des notebooks Jupyter via sa plateforme [Google Colab](https://colab.research.google.com/).
Bien que le design soit légèrement différent de celui des notebooks Jupyter classique, [Google Colab](https://colab.research.google.com/) offre néanmoins les mêmes possibilités.
Il est possible notamment de charger un notebook Jupyter que vous auriez reçu par ailleurs directement sur [Google Colab](https://colab.research.google.com/).
Enfin, cette alternative ne nécessite aucune installation puisque `Python` et l'ensemble des librairies dont vous pourriez avoir besoin sont disponibles par défaut sur la plateforme.

---

**Ressources en ligne** &emsp; L'objectif de ce cours est de vous familiariser avec `Python`.
Pour cela, chaque TP comportera un certain nombre d'exercices à faire pendant la séance.
Ces exercices ne vous présenteront néanmoins que les éléments les plus basiques et une certaine dose de travail personnel de votre part sera alors nécessaire.
De nombreuses ressources sont disponibles en ligne afin d'aller plus loin.
Vous trouverez ci-dessous une liste non-exhaustive.

| Lien | Decription |
|:---:|:---:|
| [Learnpython.org](https://www.learnpython.org/) | Site en anglais proposant une large gamme de tutoriels interactifs pour se familiariser avec `Python`. En plus des tutoriels d'introduction à `Python`, il propose également des exercices pour se familiariser avec `NumPy` et `Pandas`, deux des librairies `Python` les plus utilisées en science des données.|
| [Google's Python Class](https://developers.google.com/edu/python/) | description |
| [Coursera](https://www.coursera.org/) | description |
| [CodeWars](https://www.codewars.com/) | description |

Il existe de nombreuses autres ressources disponibles, notamment en Français.
Si l'anglais n'est donc pas votre tasse de thé, une petite recherche Google devrait vous aider.
Je vous déconseille cependant d'utiliser les outils de traduction automatique offerts par Google puisque ces derniers pourraient traduire en français des bouts de code qui ne pourraient alors plus fonctionner.

---

**Organisation du cours** &emsp; Ce cours d'introduction à `Python` se divise en 5 séances de 3h chacune.
Chaque séance abordera un aspect particulier de `Python` et se divisera de la façon suivante : environ 1h de présentation/cours/exemples suivie d'environ 2h de mise en pratique à l'aide de TP sous forme de notebooks Jupyter.
Les différentes séances aborderont les éléments suivants :
- *TP n°1 - Introduction à Python :* Le présent TP. Son objectif est de vous familiariser avec la syntaxe en `Python`.
- *TP n°2 - Introduction à NumPy :* `NumPy` est LA bibliothèque pour manipuler des vecteurs et des tableaux/matrices formant la base du calcul scientifique.
- *TP n°3 - Introduction à Matplotlib :* `Matplotlib` est la bibliothèque `Python` la plus populaire et la plus complète pour tracer des figures, éléments essentiels pour présenter les résultats de vos analyses.
- *TP n°4 - Bonnes pratiques en NumPy :* Dans ce TP, nous reviendrons à `NumPy` et discuterons en particulier les bonnes pratiques de programmation permettant de faire en sorte que votre code tourne rapidement.
- *TP n°5 - Introduction à SymPy :* `SymPy` est une bibliothèque `Python` permettant de faire du calcul symbolique de façon similaire à `Mathematica` ou `Mapple`.

Bien qu'il sera nécessaire de rendre vos notebooks Jupyter à la fin de chaque séance (ou quelques jours après), **uniquement les TP n° 3 et 4 feront l'objet d'une évaluation notée**.
Cette évaluation sera complétée par un examen dont la date reste à déterminer avec la scolarité en fonction de vos plannings.
Afin de maximiser vos chances de réussite, votre note finale sera calculée de la façon suivante

$$
\textrm{Note finale} = \max \left( \dfrac{\textrm{TP 3 + TP 4 + Examen}}{3},\ \textrm{Examen} \right).
$$

Si malheureusement la situation sanitaire ne permet pas de faire un examen en présentiel, les modalités d'évaluation seraient à re-discuter.

---
---

## Introduction à Python

`Python` est un langage de programmation qui peut s'utiliser dans un champ très large d'applications, notamment grâce à des bibliothèques (ou *packages*) spécialisées.
Il est actuellement de plus en plus utilisé dans le domaine du développement web ou pour des applications mobiles.
Il est également extrêmement répandu dans le monde scientifique, en particulier grâce aux nombreuses bibliothèques optimisées destinées au calcul numérique.
A titre d'exemple, la plupart des bibliothèques les plus utilisées aujourd'hui en *Machine Learning* et *Intelligence Artificielle* sont développées directement en `Python`.

`Python` a été conçu pour être un langage lisible et visuellement épuré.
Ainsi, il possède beaucoup moins de constructions syntaxiques que d'autres langages (par exemple les accolades { et } en C et C++).
Cette syntaxe relativement simple à prendre en main est l'une des raisons pour lesquelles `Python` est souvent le premier langage de programmation enseigné aux étudiants.
C'est maintenant à votre tour!

### Hello World!

Afficher à l'écran la chaîne de caractères `Hello World!` est traditionnellement la première chose demandée lorsque l'on apprend un nouveau langage de programmation.
Cela permet de s'assurer bien souvent que les éléments de base nécessaires à l'exécution d'un programme sont installés correctement.
Une page [Wikipedia](https://fr.wikipedia.org/wiki/Liste_de_programmes_Hello_world) dédiée vous présente comment faire dans plus de 150 langages de programmation différents.
Alors que certains langages ne nécessitent qu'une seule ligne de code pour réussir ce tour de force, d'autres (souvent plus anciens) peuvent avoir besoin d'une dizaine de lignes ou plus!

En `Python`, la ligne de code suivante suffit

```
print("Hello World!")
```

`print` est ce que l'on appelle une fonction `Python` permettant d'afficher quelque chose.
Dans le cas présent, elle prend en argument d'entrée une chaîne de caractères (délimitée par `" "`) et l'affiche à l'écran.
A votre tour maintenant de vérifier que tout fonctionne correctement en affichant `Hello World!` à l'écran.
Une fois que vous avez entré votre ligne de code, vous pouvez l'exécuter en choississant la cellule et en appuyant sur `Exécuter` dans la barre de menu ou alors directement à l'aide des touches `Shift + Entrée`.

<!-- grade_id: cell-c20018022f79b200, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Si tout se passe bien, vous devriez voir les mots `Hello World!` affichés sous la cellule Jupyter ci-dessus.
Si ce n'est pas le cas, cela signifie qu'il y'a un problème.
Contactez-moi le cas échéant.

Il est possible en `Python` d'ajouter des commentaires à votre code afin de documenter ce que vous faites.
Peu importe ce que vous écrivez comme commentaires, ces derniers ne seront pas exécutés par le langage.
**L'usage de commentaires est fortement recommandé** pour deux raisons principales :
1. *Pour vous* : les commentaires vous peuvent être utilisés comme des penses-bêtes pour vous souvenir de différentes choses dans votre code, notemment lorsque vous le mettez de côté un moment et recommencer à travailler dessus quelques jours ou semaines plus tard.
2. *Pour les autres* : il arrive bien souvent que le code que vous écriviez soit amené à être partager avec d'autres personnes (vos collègues, clients, ami.e.s, etc). Les commentaires servent alors à documenter votre code facilitant alors sa compréhension par quelqu'un d'autre, en particulier si le code est long et compliqué ou si vous n'êtes pas à côté.

En Python, un commentaire commence par le caractère `#`.
Tout ce qui suit, tant que cela reste sur la même ligne, sera considéré comment un commentaire par l'interpréteur `Python`.
Ceci est illustré ci-dessous.

```python
# Commentaire : print("Hello World!") affiche la chaîne de 
# caractère à l'écran.
print("Hello World!")
```

Si vous avez exécuté la cellule ci-dessus, vous devriez voir que `Python` n'a absolument pas pris en compte ce qui était écrit dans les commentaires et n'a exécuté que le code nécessaire à afficher `Hello World!` à l'écran.

<div class="alert alert-block alert-info">
    <b>Exercice :</b> Utilisez la fonction <code>print</code> pour afficher la phrase que vous voulez à l'écran.
</div>

<!-- grade_id: cell-3aa7f4402d574d6f, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---

### Variables et types

Une **variable** est un élement qui associe un *nom* (l'identifiant) à une *donnée* (ou *valeur*) qui sera implantée dans la mémoire de l'ordinateur.
Ainsi, la ligne suivante

```
msg = "Hello World!"
```

crée une variable appelée `msg` dans laquelle on stocke la chaîne de caractères `Hello World!`.
A noter que la valeur de la variable peut varier au cours de l'exécution du programme.

<div class="alert alert-block alert-info">
    <b>Exercice :</b> Expliquez avec vos mots ce que fait le programme suivant.

    msg = "Hello World!"

    print(msg)

    msg = "Goodbye World!"

    print(msg)
</div>

<!-- grade_id: cell-7c8b40d5cfd7293e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Chaque variable créée est d'un **type** donné.
Dans l'exemple précédent, la variable `msg` est une chaîne de caractère, ou *string* en anglais (notée `str` en `Python`).
Il existe de nombreux types différents.
En ce qui nous concerne, les types principaux avec lesquels nous seront amenés à travailler sont les suivants :
- `str` : chaîne de caractères (même si elle ne contient qu'un seul caractère).
- `int` : nombre entier (e.g. 1, 2, 3, 4 etc).
- `float` : nombre à virgule (e.g. 3.1415, 10000.838, etc).

&#x26A0;&#x26A0;&#x26A0; En `Python`, les variables `x = 1` et `y = 1.0` sont de types différents !
Bien qu'ils soit mathématiquement équivalents, `x` est enregistré en mémoire comme étant un entier (`int`) tandis que `y` est un flottant (`float`).

Afin de connaître le type d'une variable, vous pouvez utilisez la fonction `type`.
Par exemple, pour connaître le type de `x` et `y` nous pourrions utiliser le code suivant

```
x = 1
y = 1.0

# --> Affiche à l'écran le type de la variable x.
print(type(x))

# --> Affiche à l'écran le type de la variable y.
print(type(y))
```

<div class="alert alert-block alert-info">
    <b>Exercice :</b> Considérez les variables suivantes

```
msg = "Je suis un entier"
x = 1234
y = 0.5
z = 3/1
```

Déterminez le type de chacune d'elle.
Avez-vous un résultat surprenant ?
Si oui, essayez de l'expliquer.
    </div>

```python

```

<!-- grade_id: cell-0031067e982a9f07, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<div class="alert alert-block alert-info">
    <b>Exercice :</b> Dans cet exercice, on se propose de créer une chaîne de caractère, un nombre entier et un nombre flottant.
    La chaîne de caractère <code>Je suis une phrase</code> devra être stockée dans une variable appelée <code>mystr</code>.
    Le nombre flottant devra être égal à <code>3.1415</code> et stocké dans une variable <code>myfloat</code>.
    Enfin, le nombre entier sera égal à <code>42</code> et stocké dans la variable <code>myint</code>.
</div>    

<!-- grade_id: cell-9b0fcf6d9e1b41b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f2fbd0819fff531, points: 0, grade: true, solution: false -->
```python

```

---

### Les Listes

Une **liste** (*list* en anglais) est une structure de données contenant une série de valeurs, possiblement toutes de type différent.
Une liste peut par ailleurs contenir autant de valeurs que vous le souhaitez (dans la limite de la mémoire disponible sur votre ordinateur) et il est possible de récupérer l'une des valeurs de la liste en utilisant sa position dans la liste (son **indice** ou *index* en anglais).
Cela en fait une structure de données extrêmement générale.
Voici quelques exemples.

```python
animaux = ["lion", "tigre", "otarie", "cheval", "koala"]
mixte = ["lion", 42, "éléphant", 3.1415]
```

Comme vous pouvez le voir, une liste est déclarée par une série de valeurs séparées par des virgules, et le tout encadré par des **crochets**.

&#x26A0;&#x26A0;&#x26A0; Si il est possible de récupérer n'importe quel élément de la liste en utilisant son indice, il est important de noter qu'en `Python` la première position dans la liste correspond à l'indice `0`!
Ainsi, pour une liste comportant `n` éléments, le premier élément est à la position `0`, le deuxième élément à la position `1`, et ainsi de suite jusqu'au n-ième élément à la position `n-1`.
Le code ci-dessous illustre comment récupérer le deuxième élément de la liste `animaux` et ensuite l'afficher à l'écran.

```
data = animaux[1]
print(data)
```

<div class="alert alert-block alert-info">
    <b>Exercice :</b> Récupérez le quatrième élément de la liste <code>mixte</code> et affichez le à l'écran.
</div>

<!-- grade_id: cell-2fe0aca5e76843ec, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Concaténer sur les listes

Tout comme pour des chaînes de caractères, l'opération consistant à joindre deux listes s'appelle la **concaténation**.
Cette opération peut être faite de façons en `Python`.
Ceci est illustré ci-dessous.

```python
# En utilisant l'opérateur +
resultat_1 = mixte + animaux
print("En utilisant l'opérateur + :", resultat_1, "\n")

# En utilisant la fonction extend.
resultat_2 = mixte.copy()    # On fait une copie de la liste mixte histoire de ne pas réécrire par-dessus.
resultat_2.extend(animaux)
print("En utilisant la fonction extend :", resultat_2)
```

Bien qu'elles donnent le même résultat, ces deux manières de faire fonctionnent de façon différente.
Lorsque l'on fait `mixte + animaux`, `Python` concatène les deux listes mais stocke le résultat dans un nouvel endroit en mémoire.
Si on veut récupérer ce résultat, il est alors nécessaire de créer une nouvelle variable (ici `resultat_1`).
La fonction `extend` se comporte différement.
Après avoir copié le contenu de la liste `mixte` dans une nouvelle variable `resultat_2`, cette dernière contient les éléments suivants

```
["lion", 42, "éléphant", 3.1415]
```

En faisant `resultat_2.extend(animaux)`, `Python` concatène les deux listes mais ne stocke pas le résultat dans un nouvel endroit mémoire.
A la place, il augmente la mémoire allouée à la variable `resultat_2` et y ajoute les éléments de la liste `animaux`.

Ces deux façons de faire correspondent donc à deux usages légèrement différents en pratique :
- Si il est nécessaire de garder en mémoire l'ancienne liste, alors préférez l'opérateur `+`.
- En revanche, si vous souhaitez juste concaténer les deux listes sans avoir besoin de garder une trace de ce que la liste initiale était, alors utilisez `extend`.

Si ce n'est pas très clair, ce n'est pas grave.
Dans le doute, utilisez `+`.

#### Tranches (slicing) et indiçage négatif

L'un des avantages des listes est la possibilité de sélectionner une partie d'une liste `x` en utilisant un indiçage du type `x[m:n]`.
Cette commande permet de récupérer tous les éléments de la liste `x` du `m`-ème au `(n-1)`-ème (i.e. de `m` inclus à `n` exclus).
On dit alors que l'on récupère une **tranche** (**slice** en anglais) de la liste.
Rappelez-vous par ailleurs que `Python` se base sur un modèle d'indice commençant à `0`!

```python
animaux = ["lion", "tigre", "otarie", "cheval", "koala"]

# --> Affiche les trois premiers éléments.
print( animaux[0:3] )
#print( animaux[:3] )     # Alternative.

# --> Affiche tous les élements à partir du 2 ème.
print( animaux[1:] )

# --> Affiche tous les éléments du 2 ème à l'avant dernier.
print( animaux[1:-1] )
```

Il est également possible de préciser le pas, i.e. le nombre d'éléments que l'on souhaite sauter à chaque fois.

```python
x = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

# --> Pas de 2 pour récupérer tous les nombres paires de la liste.
print( x[::2] )

# --> Récupère tous les éléments, du 1er au 9eme par pas de trois.
print( x[0:9:3] )
```

Le modèel d'indiçage est donc le suivant : `[début:fin:pas]`.
Notez q'un pas de `-1` inverse la liste, e.g.

```python
# --> Inverse la liste.
print( animaux[::-1] )
print( x[::-1] )
```

L'usage d'un indice négatif est également très utile pour récupérer les derniers éléments d'une liste, peu importe sa taille.
Ainsi, `x[-1]` récupère le dernier élément tandis que `x[-2]` récupère l'avant-dernier, etc.

#### Les fonctions `len()`, `range()` et `list()`

La fonction `len(x)` vous permet de connaître la longueur de la liste `x`, c'est à dire le nombre d'éléments qu'elle contient.

```python
x = [0, 1, 2, 3, 4, 5]
print("Longueur de la liste : ", len(x) )
```

Il existe deux autres fonctions qui peuvent vous être très utiles : `range()` et `list()`.
Vous pouvez regarder leur documentation et essayer de comprendre ce qu'elles font en tapant `range?` et `list?` dans la cellule ci-dessous.

```python
range?
```

```python
list?
```

#### Listes de listes

Une liste peut ne pas contenir que des nombres ou des chaines de caractères mais également d'autres listes!
Ainsi, l'exemple ci-dessous (inspiré d'[ici](https://python.sdv.univ-paris-diderot.fr/04_listes/)) est parfaitment fonctionnel.

```python
# --> Liste des enclos.
enclos_1 = ["girafe", 4]
enclos_2 = ["tigre", 2]
enclos_3 = ["singe", 5]

# --> Zoo.
zoo = [enclos_1, enclos_2, enclos_3 ]
print(zoo)
```

Pour accéder à un élément de la liste `zoo`, on utilise l'indiçage habituel, e.g.

```python
zoo[1]
```

Il est également possible d'accéder à un élément de la sous-liste en utilisant un double indiçage

```python
zoo[1][0]
```

En absence de librairie spécialisée, les listes de listes peuvent être utilisées pour créer des matrices.
Ainsi, la matrice

$$
\mathbf{A}
=
\begin{bmatrix}
1 & 2 & 3 \\
4 & 5 &  6 \\
7 & 8 & 9
\end{bmatrix}.
$$

peut être créée de la façon suivante

```python
ligne_1 = [1, 2, 3]
ligne_2 = [4, 5, 6]
ligne_3 = [7, 8, 9]

A = [ligne_1, ligne_2, ligne_3]

print(A)
```

Ainsi, chaque ligne est stockée dans une liste.
Toutes ces listes sont ensuite stockées elles-même dans une liste qui formera la matrice.
Nous verrons un peu plus tard comment faire un produit matrice-vecteur ou matrice-matrice en pur `Python` en utilisant ces listes de listes.

---

### Opérations élémentaires

Intéressons nous maintenant aux opérations élémentaires que l'on peut appliquer aux différents types (`str`, `int`, et `float`) ainsi qu'aux `list`.

#### Opérateurs arithmétiques

De base, `Python` ne possède que relativement peu d'opérations arithmétiques élémentaires.
Vous trouverez ci-dessous une suite d'exemples sur comment manipuler les nombres entiers (`int`) et ceux à virgules (`float`).

```python
# --> Additioner deux nombres.
a, b = 1, 2
c = a + b
print(c)
```

```python
# --> Soustraire deux nombres.
a, b = 1, 2
c = a - b
print(c)
```

```python
# --> Elever un nombre à une puissance.
a = 2**2
print(a)

b = 2**1.5
print(b)
```

```python
# --> Diviser un nombre
a = 4
b = a / 2
print(b, type(b))
```

&#x26A0;&#x26A0;&#x26A0; Dans l'exemple ci-dessous, la division `/` retourne un `float` et non un entier!
Si vous souhaitez effectuer une division entière, il faut alors utiliser l'opérateur `//`.

```python
# --> Division entière.
a = 4
b = a //4
print(b, type(b))
```

```python
# --> Calcul du reste d'une division.
a, b = 11, 3
c = a % b
print(c, type(c))
print()

# --> Calcul du reste d'une division.
a, b = 11.5, 3
c = a % b
print(c, type(c))
```

```python
# --> Retourne le maximum entre deux (ou plusieurs) nombres.
print("Max(4, 10) : ", max(4, 10))

# --> Retourne le minimum entre deux (plusieurs) nombres.
print("Min(4, 10) :", min(4, 10, 2, -4))
```

#### Opérations sur les chaînes de caractères

Les opérateurs `+` (concaténation) et `*` (duplication) peuvent être utilisés sur des chaînes de caractères.
Les fonctions `len`, `max` et `min` peuvent également être utilisées pour retourner la longueur de la chaîne de caractères, le premier (ou dernier) caractère dans l'alphabet apparaissant dans la chaîne.

```python
# --> Opérateur + pour des chaînes de caractères : concaténation.
str_1, str_2 = "abc", "def"
print(str_1 + str_2)
```

```python
# --> Opérateur * pour des chaînes de caractères : duplication.
str_1 = "abc"
print( 3*str_1 )
```

```python
# --> Fonction len : longueur de la chaîne de caractères.
string = "abcdef"
print( len(string) )
```

```python
# --> Fonction max :
string = "Je suis une phrase avec des chiffres 1 2 3 4."
print( max(string) )
```

```python
# --> Fonction min :
string = "kdpdifcsq"
print("Min sans espace :", min(string) )

# --> Fonction min : avec espace.
string = "Je suis une phrase."
print("Min avec espace :", min(string) )
```

Ces quelques exemples sur les chaînes de caractères sont purement illustratifs.
Nous ne nous attarderons néanmoins pas trop dessus, notre objectif étant de faire du calcul scientifique avec `Python` et non du traitement de texte.
Si malgré vous souhaitez en savoir plus, vous pouvez vous référrer au tutoriel suivant : [learnpython.org : Basic string operations](https://www.learnpython.org/en/Basic_String_Operations).

#### Opérations sur les listes

Tout comme pour les chaînes de caractères, les opérateurs `+` et `*` peuvent être utilisés pour concaténer et dupliquer des listes comme illustré ci-dessous.

```python
# --> Concaténation de listes.
list_1 = [1, 2, 3]
list_2 = ["a", "b", "c"]

print( list_1 + list_2 )
```

```python
# --> Duplication de listes.
list_1 = [1, "a", 4]

print( 3*list_1)
```

D'autres fonctions peuvent également être appliquées aux listes.
Dans la suite, nous nous intéresserons uniquement à des listes de nombres.

```python
# --> Sommes des éléments d'une liste.
list_1 = list(range(1, 10))
print( list_1 )
print("Somme :", sum(list_1))
```

```python
# --> Maximum d'une liste.
list_1 = [1, 2, 4, -7, 10, -128]
print("Maximum : ", max(list_1) )

# --> Minimum d'une liste.
print("Minimum : ", min(list_1) )
```

```python
# --> Trier une liste par ordre croissant.
list_1 = [48, 14, 5, -3, 1, 128]
print("Liste initiale :", list_1 )

list_1.sort()
print("Liste triée :", list_1)
```

&#x26A0;&#x26A0;&#x26A0; L'application de `sort` trie la liste `in-place`, c'est à dire qu'elle ré-écrit la liste triée dans le même espace mémoire.
L'agencement initial de la liste est alors perdu.
Si l'on veut garder la liste initiale, il faut alors utiliser la fonction `sorted`.

```python
list_1 = [48, 14, 5, -3, 1, 128]
list_2 = sorted(list_1)

print("Liste 1 : ", list_1)
print("Liste 2 : ", list_2)
```

---

### Les conditions

Il est par moment nécessaire de vérifier certaines conditions avant d'effectuer une opération particulière, ou alors choisir une certaine opération en fonction des conditions qui ont été vérifiées.
Ceci s'effectue grâce à l'algèbre de Bool et aux opérateurs `>`, `<`, `<=`, `>=`, `==`, `and` et `or`.
Quelques exemples sont proposés ci-dessous.

```python
x = 2

# --> Condition d'égalité
print("Est-ce que x est égal à 2 ?", x == 2, "\n")

# --> Inférieur et inférieur ou égal.
print("Est-ce que x est inférieur à 2 ?", x < 2, "\n")

print("Est-ce que x est inférieur ou égal à 2 ?", x <= 2)
```

Il est également possible de "chaîner" ces opérateurs.
Ainsi, si l'on souhaite savoir si `x` est compris entre 1 et 4, on utilisera le code suivant.

```python
x = 2

print("Est-ce que x est compris entre 1 et 4 ?", 1 <= x <= 4, "\n")

print("Est-ce que x est compris entre -10 et -2 ?", -10 <= x <= -2)
```

Il est possible de choisir quelle opération effectuée en fonction des conditions qui ont été (ou non) vérifiées.
Ainsi, l'exemple ci-dessous affiche la valeur de `x` si ce dernier est un nombre impair, sinon il affiche que le nombre est pair.

```python
x = 11

if x % 2 == 1 and type(x) == int: # si x est impair (x % 2 retourne le reste de la division par 2) ET si c'est un entier.
    print("La valeur de x est égale à ", x)
else:
    print("x est un nombre de pair ou un nombre à virgule.")
```

Nous ne rentrerons pas dans tous les détails des opérations possibles à l'aide de l'algèbre de Bool.
Pour plus de détails, il existe de très nombreux exemples en ligne.

---

### Les boucles `for` et `while`

```python
for i in range(10):
    print(i)
```

```python
i = 10000
while i > 10:
    print(i)
    i = i / 2
```

#### Compréhension de listes

```python
# --> Construit une liste avec tous les carrés des entiers de 1 à 10.
list_1 = [i**2 for i in range(1, 11)]
print(list_1)
```

```python
# --> Construit une liste ne comportant que les nombres pairs de 1 à 50.
list_1 = [i for i in range(1, 51) if i % 2 == 0]
print(list_1)
```

---

### Comment définir des fonctions en `Python` ?

Une fonction est un petit programme contenant uen suite d'instructions que l'on peut appeler avec un nom.
Par exemple, dans les exemples précédents, `max`, `len` ou `min` sont des fonctions.
Elle prends (éventuellement) en entrée des **arguments d'entrée** et retourne léventuellement) le résultat de ces instructions comme **arguments de sortie**.
Trois exemples de telles fonctions vous sont présentés ci-dessous.

**Fonction ne retournant aucune valeur**

```
def nom_de_la_fonction(argument 1, argument_2, ...):
    Vos opération
```

```python
# --> Definition de la fonction.
def affiche_info_personnelles(nom, prenom, age):
    
    # --> Tests pour vérifier la validité des arguments d'entrée.
    assert type(nom) == type(prenom) == str
    assert type(age) == int
    
    # --> Affiche à l'écran vos informations.
    print("Vous vous appellez {nom} {prenom} et vous avez {age} ans.".format(nom=nom, prenom=prenom, age=age))
    
# --> Utilisation de la fonction.
affiche_info_personnelles("Dupont", "Alice", 24)
affiche_info_personnelles("Martin", "Bob", 42)
```

**Fonction retournant quelque chose**

```
def nom_de_la_fonction(argument_1, argument_2, ...):
    Vos opération
    return variable_de_retour_1, variable_de_retour_2, ...
```

```python
# --> Définition de la fonction.
def my_add(a, b):
    c = a + b
    return c

# --> Définition de la fonction.
def my_add_and_my_sub(a, b):
    c = a + b
    d = a - b
    return c, d
```

**Fonction lambda :** Définition d'une fonction qui tient sur une seule ligne.

```python
# --> Fonction lambda (définition d'une fonction en une seule ligne)
my_add_bis = lambda x, y : x + y
```

---
---
<div class="alert alert-block alert-danger">
    <h2>Exercices notés</h2>
</div>

Vous trouverez ci-dessous une liste d'exercices.
Il est peu probable que vous ayez le temps de les faire pendant la séance.
Il s'agit donc essentiellement d'exercices à faire chez vous.
Attention, même si ils ne rentrent pas en compte dans votre note finale, ces exercices seront notés!
Cette note de TP sera utilisée pour juger en votre faveur ou défaveur si jamais il venait à vous manquer quelques points pour valider.
**Ces exercices sont également des exemples de ce qui pourrait vous attendre le jour de l'examen.**
A vous de voir à quel point vous êtes joueur/joueuse...

**Exercice n°1 :** Ecrivez une fonction retournant combien il y a de nombres pairs entre deux entiers `a` et `b` ainsi que l'ensemble de ces nombres pairs.
Attention, si `a > b` en entrée de votre fonction, alors vous devez les intervertir!

<!-- grade_id: cell-4e0781c757e33223, points: 0, grade: false, solution: true -->
```python
def nbres_pairs(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7e4f46c416770fbb, points: 1, grade: true, solution: false -->
```python
assert nbres_pairs(0, 2) == ([0, 2], 2)
assert nbres_pairs(4, 3) == nbres_pairs(3, 4) ==  ([4], 1)

```

**Exercice n°2 :** La suite de Fibonacci est donnée par la récurrence suivante

$$
\begin{aligned}
f_0 & = 0 \\
f_1 & = 1 \\
f_n & = f_{n-1} + f_{n-2} \quad \forall n \geq 2.
\end{aligned}
$$

Ecrivez une fonction retournant le `n`-ème nombre de Fibonnacci.

<!-- grade_id: cell-f8cd9b5ba9dd637e, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0c46897007985164, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(10) == 55
```

**Exercice n°3 :** Ecrivez une fonction `Python` calculant la somme des `n` premiers entiers au carré, i.e.

$$
s = \sum_{i=1}^n i^2
$$

<!-- grade_id: cell-12763e0cab88430e, points: 0, grade: false, solution: true -->
```python
def sum_square(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-54774498d26faa11, points: 1, grade: true, solution: false -->
```python
assert sum_square(1) == 1
assert sum_square(10) == 385
assert sum_square(100) == 338350
```

**Exercice n°4 :** Ecrivez une fonction `Python` calculant la somme des `n` premiers nombres impairs au carré, i.e.

$$
s = 1^2 + 3^2 + 5^3 + \cdots
$$

<!-- grade_id: cell-721029d603ea41d8, points: 0, grade: false, solution: true -->
```python
def impairs_sum_square(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a45d878b14e5a726, points: 1, grade: true, solution: false -->
```python
assert impairs_sum_square(1) == 1
assert impairs_sum_square(10) == 1330
assert impairs_sum_square(100) == 1333300
```

**Exercice n°5 :** Ecrivez une fonction `Python` calculant $n! = n \times (n-1) \times (n-2) \times \cdots \times 3 \times 2 \times 1$. Par définition, $0! = 1$ et votre fonction doit donc gérer ce cas particulier également.

<!-- grade_id: cell-9395a8408b75a034, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1b13cb12cf250334, points: 1, grade: true, solution: false -->
```python
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(2) == 2
assert factorial(4) == 24
assert factorial(8) == 40320
assert factorial(16) == 20922789888000
```

**Exercice n°6 :** En utilisant les listes de listes, écrivez une fonction faisant le produit d'une matrice et d'un vecteur.
Votre fonction doit fonctionner peu importe le nombre de lignes et de colonnes de la matrice.

Pour rappel, le produit matrice-vecteur peut s'écrire de la façon suivante

$$
b_i = \sum_{j=1}^n A_{ij} x_j.
$$

Vous aurez besoin d'utiliser les boucles `for` pour implémenter cette somme.

<!-- grade_id: cell-5e8d4e83130a77a4, points: 0, grade: false, solution: true -->
```python
def matvec(A, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f428d6efc819842b, points: 2, grade: true, solution: false -->
```python
# --> Test n°1
A, x = [[1, 2], [3, 4]], [1, 1]
assert matvec(A, x) == [3, 7]

# --> Test n°2
A, x = [[1, 2], [3, 4], [5, 6]], [1, 1]
assert matvec(A, x) == [3, 7, 11]
```

**Exercice n°7 :** Ecrivez maintenant une fonction `Python` calculant le produit matrice-matrice.
Pour rappel, les entrées $C_{ij}$ de la matrice $\mathbf{C} = \mathbf{AB}$ sont données par

$$
C_{ij} = \sum_{k=1}^n A_{ik} B_{kj}
$$

où $n$ est le nombre de colonnes de la matrice $\mathbf{A}$ mais aussi le nombre de lignes de la matrice $\mathbf{B}$.

<!-- grade_id: cell-7a040c01d58c1841, points: 0, grade: false, solution: true -->
```python
def matmat(A, B):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-42365c04af4d5a85, points: 2, grade: true, solution: false -->
```python
# --> Test n°1
A = [[1, 2], [3, 4]]
I = [[1, 0], [0, 1]]
assert matmat(A, I) == matmat(I, A) == A

# --> Test n°2
A = [[1, 2], [3, 4], [5, 6]]
B = [[0, 1], [1, 0]]
assert matmat(A, B) == [[2, 1], [4, 3], [6, 5]]

# --> Test n°3
A = [[1, 2], [3, 4]]
B = [[1, 2], [3, 4]]
assert matmat(A, B) == [[7, 10], [15, 22]]
```

**Exercice n°8 :** Ecrivez une fonction affichant à l'écran la table de multiplication pour un entier `n` (on supposera que `n` est compris entre 1 et 10).
Pour en savoir plus sur comment manipuler des chaînes de caractères, vous pouvez regarder le tutoriel suivant : [learnpython.org : String formatting](https://www.learnpython.org/en/String_Formatting).

<!-- grade_id: cell-675fce58192daecf, points: 1, grade: true, solution: true -->
```python
def table_multiplication(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```